## Assignment 2 - NLP using Rules
_Solutions have to be submitted by Monday, May 21, 23:59 (UTC+1)._

---

### Exercise 1: Finite-State Transducers (1 point)

**(a) Consider the task of automatic spelling correction. Do you think a Finite State Transducer (FST) would be needed for this task or would a Finite State Automaton (FSA) suffice? Shortly explain your choice in one sentence.**

<!-- grade_id: cell-37695585c9f93b0e, points: 0.5, grade: true, solution: true -->
An FST would be needed, because we need to map input and output. If a word is typed wrongly as input, the FST can correct the wrong input word, and swaps the wrong input word with an correct output word. 

**(b) Can every task that is solvable with an FSA also be solved with an FST? Shortly explain your answer in one sentence.**

<!-- grade_id: cell-f78b09c1d51e73a9, points: 0.5, grade: true, solution: true -->
Yes, every task that is solvable with an FSA is solvable with an FST, because an FST is an extension of an FSA. FSA can only read on strings, while FSTs can read and write them. So FSTs can read strings, just like an FSA.

### Exercise 2: Lexicon-based approaches (2 points)

**(a) Consider creating a lexicon of companies participating in the New York Stock Exchange. The lexicon should not only list companies, but also contain useful information about each, such as the name of the CEO. From the types of lexicons presented in the lecture part on "NLP using lexicons", which one is the best option? Explain shortly.**

<!-- grade_id: cell-e42a9292be875ce3, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) Given the following three NLP tasks below, write down if they could be effectively approached with a lexicon-based method. Explain your answers shortly (no point without explanation).**

1. **Finding all named entities (disregarding what entity type they have) in a well-formed English text.**
2. **Extracting natural language arguments from a corpus of debate speeches.**
3. **Finding sentences that contain explicit insults.**

<!-- grade_id: cell-b869430672f07529, points: 1.5, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 3: Decision trees (3 points)

**Tokenization of English texts faces problems similar to those of sentence splitting presented in the lecture. Up to a certain accuracy, however, it can also be approached using a decision tree.**

**(a) Create and draw a graphical representation of a decision tree that is able to tokenize a given English text in a similar manner to the decision tree for sentence splitting presented in the lecture. The decision tree should be able to deal at least with the following cases (in that order):**

1. **Split at whitespaces and sentence endings.**
2. **Split at a comma, but do not split inside numbers.**
3. **Split at hyphens that are not part of a multi-word phrase.**
4. **Split at full stops, but not if the dot is part of a number, abbreviation or URL.**

**For simplicity, you can expect functions to be available that detect sentence endings, numbers, multi-word phrases, as well as abbreviations and URLs. You might use the notation introduced in the lecture for the sentence splitter for the graphical representation of you decision tree.**

_Note: You can include images in the markdown cell below by uploading an image file to the same directory as this notebook and then adding it in markdown via `![alt-text](image-file-name.png)`. You are free to choose between drawing the tree by hand (and then scanning it) or any digital drawing program of your liking. Make sure, however, that all texts and connections are clearly readable. If this is not the case, you will not get any points for this task._

<!-- grade_id: cell-04ba952973083ee1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**(b) What potential problems do you see with tackeling tokenization using decision trees? Shortly explain your answer and come up with at least two edge-cases that are not covered by the rules stated in (a) for the problem of tokenization of a well-formed English text.**

<!-- grade_id: cell-5276b0a65ca3ee47, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 4: Text stemming (1 point)

**Similar to the Porter Stemmer, the Krovetz Stemmer uses a set of rules to determine the stem of a word. In addition, it introduces the use of lexicons to cover most of the well-known edge cases, in which teh Porter Stemmer produces bad results.**

**The general process is as follows:**

1. **If input token present in lexicon, replace with stem.**
2. **If not present, check token for choppable inflection suffixes.**
3. **If chopped token present in lexicon, replace with stem.**
4. **If still not present, try to add different suffixes.**

**What do you think might be advantages and disadvantages of this approach when compared to the Porter Stemmer? Name at least one each and shortly explain your answer.**

<!-- grade_id: cell-5628d4c294e127ae, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 5: Template-based generation (3 points)

**Consider the task of template-based natural language generation in the domain of sports event reporting.**

**For a kind of sports of your choice, come up with two template sentences that has to be filled with texts from data. Each template should have at least 3 slots to be filled. For each of those cases, list and explain the following things:**

1. **The intended goal of the template and what it should convey.**
2. **The data (i.e. the entities, attribute and relations) that would be required to generate the sentence and the knowledge to order the retrieved data in the planning phase.**
3. **At least one rule that are in applied in each, the Lexicalization and Referring expression generation step and why it makes sense to apply this rule.**

**For each template, also provide an example.**


_Note: The data and rules need to be different from those presented in the lecture._

<!-- grade_id: cell-6302a4abd6bd1700, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

---

### Exercise 6: Programming task 1 (5 points)

**In this task you shall implement a rule-based sentence splitter.**

**Write a function `my_sentence_splitter` that takes as an input a text of type string, parse the string character by character, and applies a set of rules in order to identify all positions where the text should be split. Then the function should return a list of strings representing the sentences extracted from the text. For examples of possible rules refer to the lecture part on rule-based approaches. You are free to add more rules that increase your performance.**

**Shortly describe your appraoch to sentence splitting and your approach in the docstring of the function.**

**Your code will be evaluated on a ground-truth sample of texts. The provided function `evaluate_sentence_splitter` will return the accuracy of your sentence splitter on the sample. Try to optimize your rules to get the highest accuracy. The same ground-truth texts would be used to evaluate your code. Part of the credits for this task will depend on the achieved performance.**

_Note: An empty input should also return an empty list._

_Note: The accuracy is computed as the ratio of the correct sentences your approach produces to the total number of ground-truth sentence._

```python
def evaluate_sentence_splitter(sentence_splitter) -> float:
    """
    This function takes as an input a function `sentence_splitter` and returns the
    accuracy of this function evaluated against a text sample.

    The function `sentence_splitter` should accept a string as a parameter and
    return a list of strings.
    """
    # Read data files
    with open("text_raw.txt", "r", encoding="utf-8") as f:
        raw_texts = f.read().split("\n")
    with open("text_sentences.txt", "r", encoding="utf-8") as f:
        sentences = f.read().split("<new_line>")
    processed_texts = list(map(lambda x: x.strip().split("\n"), sentences))
    
    scores = []
    # Create sentence pairs from raw and processed texts
    for raw_text, text_sentences in zip(raw_texts, processed_texts):
        # Split text into sentences with provided splitter
        pred_sentences = sentence_splitter(raw_text)
        
        # Test for the amount of intersection between the two sentences
        if len(pred_sentences) == 0:
            score = 0
        else:
            score = len(set(pred_sentences).intersection(set(text_sentences)))/len(set(text_sentences))
        
        scores.append(score)

    # Return the accuracy
    return round(sum(scores)/len(scores), 3)
```

<!-- grade_id: cell-dfc295e9b48904ff, points: 0, grade: false, solution: true -->
```python
def my_sentence_splitter(text: str) -> list[str]:
    """Describe implementation details here..."""
    # YOUR CODE HERE
    raise NotImplementedError
```

```python
accuracy = evaluate_sentence_splitter(my_sentence_splitter)
print(f"Accuracy of your sentence splitter is: {accuracy}")
```

<!-- grade_id: cell-d2024eb9eec10d48, points: 5, grade: true, solution: false -->
```python
########## BEGIN OF AUTOMATED TESTS
# ----------------------------------------------------
# Tests for the `my_sentence_splitter` function
texts = [
    "text1 contains two sentences. This is the second sentence.",
    "text2 contains this only one sentence."
]
texts_sentences = [my_sentence_splitter(text) for text in texts]
assert type(texts_sentences) == list
assert len(texts_sentences[0]) == 2
assert len(texts_sentences[1]) == 1
########## END OF AUTOMATED TESTS
```

### Exercise 7: Programming task 2 (5 points)

**In this task you shall implement a rule-based morphological analyzer as a Finite-State Transducer (FST) for the task of English Nominal Number Inflection. Use the FST presented in the lecture on rule-based NLP.**

**(a) Implement the function `nominal_number_inflection` that takes as an input a string and returns the nominal number inflection of it as a list of strings (e.g: `["cat", "N", "Sg"]`) if it can be parsed using the defined state automata or `None` otherwise. The function should process the string character by character and try to build the inflection version of the noun. For more details on the nominal number inflection, refer to the slides of the lecture part on rules.**

**We provide sample lexicons for regular nouns of type 1 and 2, as well as for both singular and plural irregular nouns. Use these lexicons in your code to search for noun matches.**

**Use the docstring of the function to shortly describe your implementation approach.**

<!-- grade_id: cell-b03f7f62a4ff316a, points: 0, grade: false, solution: true -->
```python
def nominal_number_inflection(text: str) -> list[str]:
    """Describe implementation details here..."""
    # YOUR CODE HERE
    raise NotImplementedError
```

<!-- grade_id: cell-81e7baeb9a8b9ad6, points: 4, grade: true, solution: false -->
```python
########## BEGIN OF AUTOMATED TESTS
# ----------------------------------------------------
# Tests for the `nominal_number_inflection` function

# Regular noun type 1 Lexicon
reg_noun_type1 = set(["cat", "dog", "head"])
# Regular noun type 1 Lexicon
reg_noun_type2 = set(["bus", "hero"])
# Irregular singular noun Lexicon
irreg_singular = set(["mouse", "person", "leaf"])
# Irregular plural noun Lexicon
irreg_plural = {"people": "person", "mice": "mouse", "leaves": "leaf"}

nouns = ["cat", "mice"]
noun_inflections = [nominal_number_inflection(noun) for noun in nouns]
assert noun_inflections[0] == ["cat", "N", "Sg"]
assert noun_inflections[1] == ["mouse", "N", "Pl"]
########## END OF AUTOMATED TESTS
```

**(b) What potential problem do you see with matching the incoming string with an FST and lexicon this way? Shortly explain and provide at least one example.**

<!-- grade_id: cell-98a84595a987ecc7, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE